# Ejercicios de apoyo

## Ejercicio 1 (2 puntos)
Escribe una función `numeros_1(limite_inferior, limite_superior)` que devuelva los números que sean divisibles por 7 pero no sean múltiplos de 5 entre dos números dados (ambos incluidos).

<!-- grade_id: cell-8ba2dfd8e9973e4c, points: 0, grade: false, solution: true -->
```python
def numeros_1(limite_inferior, limite_superior):
    l =[]
    for i in range(limite_inferior, limite_superior+1):
        if i % 7 == 0 and i % 5 != 0:
            l.append(i)
    return l
```

<!-- grade_id: cell-74c2ff470b100b37, points: 2, grade: true, solution: false -->
```python
assert(numeros_1(0, 7) == [7])
assert(numeros_1(0, 14) == [7, 14])
assert(numeros_1(120, 350) == [126,
 133,
 147,
 154,
 161,
 168,
 182,
 189,
 196,
 203,
 217,
 224,
 231,
 238,
 252,
 259,
 266,
 273,
 287,
 294,
 301,
 308,
 322,
 329,
 336,
 343])
```

## Ejercicio 2 (2 puntos)

Escribe una función `suma(lista)` que devuelva la suma de todos los elementos de una lista.

<!-- grade_id: cell-dafe590def2b656f, points: 0, grade: false, solution: true -->
```python
def suma(lista):
    acc = 0
    for e in lista:
        acc += e
    return acc
```

<!-- grade_id: cell-0603fcd7515919ea, points: 2, grade: true, solution: false -->
```python
assert(suma([1,2,3,4,5]) == 15)
assert(suma([1,2,3,4,5,6,7]) == 28)
```

## Ejercicio 3 (3 puntos)
La función `cadena.replace(a, b)` de las cadenas reemplaza apariciones de la cadena "a" por la cadena "b" en "cadena".

¿Puedes implementar tú una función similar `reemplaza(cadena, a, b)` que devuelva otra cadena en la que se hayan reemplazado todas las apariciones de "a" en "cadena" por "b"?

<!-- grade_id: cell-1f3a4f692d246f99, points: 0, grade: false, solution: true -->
```python
def reemplaza(cadena, a, b):
    i = 0
    cadena_resultado = ''
    while i < len(cadena):
        if cadena[i:i+len(a)] == a:
            cadena_resultado += b
            i = i + len(a)
            continue
        cadena_resultado += cadena[i]
        i = i + 1
    return cadena_resultado
```

<!-- grade_id: cell-c20485dd20193354, points: 3, grade: true, solution: false -->
```python
assert(reemplaza("Una frase cualquiera", "a", "e") == 'Une frese cuelquiere')
assert(reemplaza("Una frase cualquiera", "Una ", "") == 'frase cualquiera')
assert(reemplaza("Una frase cualquiera", "Una", "No es una") == 'No es una frase cualquiera')
```

# Ejercicio 4 (4 puntos):
Escribe la función `agrupa(tupla, numero)` que recibe una tupla de elementos y devuelve otra tupla que contiene los mismos elementos en grupos del número que se proporciona por parámetro.

<!-- grade_id: cell-2dce46e1e6944945, points: 0, grade: false, solution: true -->
```python
# 
def agrupa(tupla, numero):
    res = [] 
    for i in range(0, len(tupla), numero):
        interm = []
        for j in range(numero):
            if i+j < len(tupla):
                interm.append(tupla[i+j])
        res.append(tuple(interm))
    return tuple(res)
```

<!-- grade_id: cell-bf902f62b2b6193f, points: 3, grade: true, solution: false -->
```python
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 2) == ((1,2),(3,4),(5,6),(7,8),(9,10)))
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 3) == ((1,2,3),(4,5,6),(7,8,9),(10,))) #Ojo, el último elemento es una tupla!
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 4) == ((1, 2, 3, 4), (5, 6, 7, 8), (9, 10)))
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 5) == ((1, 2, 3, 4, 5), (6, 7, 8, 9, 10)))
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 6) == ((1, 2, 3, 4, 5, 6), (7, 8, 9, 10)))
assert(agrupa((1,2,3,4,5,6,7,8,9,10), 7) == ((1, 2, 3, 4, 5, 6, 7), (8, 9, 10)))
```